<font size="4" style="color:red;"> **IMPORTANT: ** When submitting this homework notebook, please modify only the cells that start with:</font>

```python
# modify this cell
```

# Conditional Probability and Baye's Rule


## Problem

There are two urns $A$ and $B$. Urn $A$ contains $r_A$ red balls and $w_A$ white balls whereas urn $B$ contains $r_B$ red balls and $w_B$ white balls. One of the urns is picked at random and then one ball is picked at random from this urn. Write a function **conditional_probability** that calculates the conditional probability that the randomly chosen ball belonged to urn $A$ given that it is white. Assume that $\frac{r_A}{w_A}\neq\frac{r_B}{w_B}$.

<font  style="color:blue"> **Code**</font>
```python
rA, wA, rB, wB = 1., 2., 2., 1.
conditional__probability(rA, wA, rB, wB) 
```

<font  style="color:magenta"> **Output**</font>
```
0.6666666666666666
```

```python
import numpy as np
# modify this cell
experiments = 1000000

def conditional__probability_1(rA, wA, rB, wB):
    urn_A = ['r'] * rA + ['w'] * wA
    urn_B = ['r'] * rB + ['w'] * wB
    
    balls = urn_A + urn_B
    
    selected = np.random.randint(len(balls), size=experiments)
    print(selected[:10])
    print(balls)
    count = 0

    for select in selected:
        if balls[select] == 'w':
            if select < len(urn_A):
                count += 1
    
    return count / experiments
    
def conditional__probability(rA, wA, rB, wB):
    urn_A = ['r'] * rA + ['w'] * wA
    urn_B = ['r'] * rB + ['w'] * wB
    
    urns = urn_A, urn_B

    chose_urns = np.random.rand(experiments)
    select = np.random.randint(rA+wA, size=experiments)
    
    count = 0
    for i, urn in enumerate(chose_urns):
        selected = ''
        if urn < 0.5:
            selected = urn_A[select[i]]
        
        if selected == 'w':
            count += 1
    
    return count / experiments
```

```python
rA, wA, rB, wB = 1, 2, 2, 1
conditional__probability_1(rA, wA, rB, wB)
```

Output:
```
[4 0 3 0 4 1 2 4 2 3]
['r', 'w', 'w', 'r', 'r', 'w']
0.334281
```

<!-- grade_id: ex1, points: 5, grade: true, solution: false -->
```python
assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)  
assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5) 

#
# AUTOGRADER TEST - DO NOT REMOVE
#

```

Output:
```
<ipython-input-16-608ab43f6d0f> in <module>()
----> 1 assert( abs(conditional__probability(2., 4., 3., 3.) -0.5714285714285715) < 10**-5)
      2 assert( abs(conditional__probability(1., 3., 5., 2.) -0.7241379310344829) < 10**-5)
      3 
      4 #
      5 # AUTOGRADER TEST - DO NOT REMOVE

<ipython-input-14-70fb63bf0ddc> in conditional__probability(rA, wA, rB, wB)
      4 
      5 def conditional__probability(rA, wA, rB, wB):
----> 6     urn_A = ['r'] * rA + ['w'] * wA
      7     urn_B = ['r'] * rB + ['w'] * wB
      8 

TypeError: can't multiply sequence by non-int of type 'float'
```

```python
conditional__probability_1( 2, 1, 6, 4)
```

Output:
```
0.076501
```

```python
np.random.randint
```